Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Baby Names Python Exercise

Source: https://developers.google.com/edu/python/

## Intro
The Social Security administration has this neat data by year of what names are most popular for babies born that year in the USA (see social security baby names).

The files for this exercise are in the `babynames` directory. Add your code to this notebook. The files `baby1990.html baby1992.html ...` contain raw html, similar to what you get visiting the above social security site. Take a look at the html and think about how you might scrape the data out of it.

# Part A

In this notebook, implement the `extract_names(filename)` function which takes the filename of a `baby1990.html` file and returns the data from the file as a single list -- the year string at the start of the list followed by the name-rank strings in alphabetical order. `['2006', 'Aaliyah 91', 'Abagail 895', 'Aaron 57', ...]`.

Rather than treat the boy and girl names separately, we'll just lump them all together. **In some years, a name appears more than once in the html, but we'll just use one number per name. Optional: make the algorithm smart about this case and choose whichever number is smaller.**

Build the program as a series of small milestones, getting each step to run/print something before trying the next step. This is the pattern used by experienced programmers -- build a series of incremental milestones, each with some output to check, rather than building the whole program in one huge step.

Printing the data you have at the end of one milestone helps you think about how to re-structure that data for the next milestone. Python is well suited to this style of incremental development. For example, first get it to the point where it extracts and prints the year. Here are some suggested milestones:

1. Extract all the text from the file and print it
2. Find and extract the year and print it
3. Extract the names and rank numbers and print them
4. Get the names data into a dict and print it
5. Build the `[year, 'name rank', ... ]` list and print it

Earlier we have had functions just print to standard out. It's more re-usable to have the function *return* the extracted data, so then the caller has the choice to print it or do something else with it. (You can still print directly from inside your functions for your little experiments during development.)


For parsing files, it's nice to include a copy of the target
text for inspiration.

Here's what the html looks like in the baby.html files:
```
...
<h3 align="center">Popularity in 1990</h3>
....
<tr align="right"><td>1</td><td>Michael</td><td>Jessica</td>
<tr align="right"><td>2</td><td>Christopher</td><td>Ashley</td>
<tr align="right"><td>3</td><td>Matthew</td><td>Brittany</td>
...
```

<!-- grade_id: extract_names, points: 0, grade: false, solution: true -->
```python
def extract_names(filename):
    """
    Given a file name for baby.html, returns a list starting with the year string
    followed by the name-rank strings in alphabetical order.
    ['2006', 'Aaliyah 91', Aaron 57', 'Abagail 895', ' ...]
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_extract_names, points: 70, grade: true, solution: false -->
```python
names2000 = extract_names('baby2000.html')

len(names2000) == 1929

assert names2000[0] == '2000'

assert names2000[1] == 'Aaliyah 211'
assert names2000[2] == 'Aaron 41'
assert names2000[652] == 'Emily 1'
assert names2000[-2] == 'Zoey 409'
assert names2000[-1] == 'Zoie 728'

```

## Part B

Suppose instead of printing the text in the notebook, we want to write files containing the text. Implement the `save_names` that does so. For each input file 'foo.html', instead of printing in the notebook, write a new file 'foo.html.summary' that contains the summary text for that file.

<!-- grade_id: save_names, points: 0, grade: false, solution: true -->
```python
def save_names(filename):
    names = extract_names(filename)

    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: test_save_names, points: 30, grade: true, solution: false -->
```python
save_names('baby2000.html')

```