Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Basic string exercises

Source: https://developers.google.com/edu/python/

### `donuts`
Given an int count of a number of donuts, return a string

of the form 'Number of donuts: `<count>`', where `<count>` is the number

passed in. However, if the count is 10 or more, then use the word 'many'

instead of the actual count.

So `donuts(5)` returns 'Number of donuts: 5'

and `donuts(23)` returns 'Number of donuts: many'

<!-- grade_id: donuts, points: 0, grade: false, solution: true -->
```python
def donuts(count):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_donuts, points: 10, grade: true, solution: false -->
```python
assert donuts(4) == 'Number of donuts: 4'
assert donuts(9) == 'Number of donuts: 9'
assert donuts(10) == 'Number of donuts: many'
assert donuts(99) == 'Number of donuts: many'
```

### `both_ends`
Given a string s, return a string made of the first 2

and the last 2 chars of the original string,

so 'spring' yields 'spng'. However, if the string length

is less than 2, return instead the empty string.


<!-- grade_id: both_ends, points: 0, grade: false, solution: true -->
```python
def both_ends(s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_both_ends, points: 10, grade: true, solution: false -->
```python
assert both_ends('spring') == 'spng'
assert both_ends('Hello') == 'Helo'
assert both_ends('a') == ''
assert both_ends('xyz') == 'xyyz'
```

### `fix_start`

Given a string s, return a string

where all occurences of its first char have

been changed to '*', except do not change

the first char itself.

e.g. 'babble' yields 'ba**le'

Assume that the string is length 1 or more.

Hint: s.replace(stra, strb) returns a version of string s

where all instances of stra have been replaced by strb.


<!-- grade_id: fix_start, points: 0, grade: false, solution: true -->
```python
def fix_start(s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_fix_start, points: 10, grade: true, solution: false -->
```python
assert fix_start('babble') == 'ba**le'
assert fix_start('aardvark') == 'a*rdv*rk'
assert fix_start('google') == 'goo*le'
assert fix_start('donut') == 'donut'
```

### `mix_up`

Given strings a and b, return a single string with a and b separated

by a space '`<a> <b>`', except swap the first 2 chars of each string.

e.g.

  'mix', pod' -> 'pox mid'
  
  'dog', 'dinner' -> 'dig donner'

Assume a and b are length 2 or more.

<!-- grade_id: mix_up, points: 0, grade: false, solution: true -->
```python
def mix_up(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_mix_up, points: 10, grade: true, solution: false -->
```python
assert mix_up('mix', 'pod') == 'pox mid'
assert mix_up('dog', 'dinner') == 'dig donner'
assert mix_up('gnash', 'sport') == 'spash gnort'
assert mix_up('pezzy', 'firm') == 'fizzy perm'
```

### Extra exercises

### `verbing`

Given a string, if its length is at least 3,

add 'ing' to its end.

Unless it already ends in 'ing', in which case

add 'ly' instead.

If the string length is less than 3, leave it unchanged.

Return the resulting string.

<!-- grade_id: verbing, points: 0, grade: false, solution: true -->
```python
def verbing(s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_verbing, points: 10, grade: true, solution: false -->
```python
assert verbing('hail') == 'hailing'
assert verbing('swiming') == 'swimingly'
assert verbing('do') == 'do'
```

### `not_bad`

Given a string, find the first appearance of the

substring 'not' and 'bad'. If the 'bad' follows

the 'not', replace the whole 'not'...'bad' substring

with 'good'.

Return the resulting string.

So 'This dinner is not that bad!' yields:

This dinner is good!

<!-- grade_id: not_bad, points: 0, grade: false, solution: true -->
```python
def not_bad(s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_not_bad, points: 10, grade: true, solution: false -->
```python
assert not_bad('This movie is not so bad') == 'This movie is good'
assert not_bad('This dinner is not that bad!') == 'This dinner is good!'
assert not_bad('This tea is not hot') == 'This tea is not hot'
assert not_bad("It's bad yet not"), "It's bad yet not"
```

### `front_back`

Consider dividing a string into two halves.

If the length is even, the front and back halves are the same length.

If the length is odd, we'll say that the extra char goes in the front half.

e.g. 'abcde', the front half is 'abc', the back half 'de'.

Given 2 strings, a and b, return a string of the form

a-front + b-front + a-back + b-back

<!-- grade_id: front_back, points: 0, grade: false, solution: true -->
```python
def front_back(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_front_back, points: 10, grade: true, solution: false -->
```python
assert front_back('abcd', 'xy') == 'abxcdy'
assert front_back('abcde', 'xyz') == 'abcxydez'
assert front_back('Kitten', 'Donut') == 'KitDontenut'
```

# Basic file exercises

### `upperize_file`

Given a file name of **textual** file, returns the file's content in uppercase.

<!-- grade_id: upperize_file, points: 0, grade: false, solution: true -->
```python
def upperize_file(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_upperize_file, points: 10, grade: true, solution: false -->
```python
print(upperize_file('Alice-Project-Gutenberg.txt')[:1000])

```

### `process_ages_file`
Given a file name that contains names and ages, return a dictionay
of the average, minimum and maximum ages in the file.

For example, given this file:
```
Bonnie - 20
Clyde - 21
```
the function returns:
`{'average': 20,5, 'max': 21, 'min': 22}`

<!-- grade_id: process_ages_files, points: 0, grade: false, solution: true -->
```python
def process_ages_file(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_process_ages_files, points: 10, grade: true, solution: false -->
```python
assert process_ages_file('names-ages.txt') == { 'average': 61.125,
                                                'max': 84,
                                                'min': 39
                                              }
```

### `save_names_ages_sorted`

Given a dictionary names2ages that maps names to ages,
and a file name, save the data in the dictionary to a file using the exact same format as the previous task.
The file's lines should be sorted by alphabetic order of the names.

<!-- grade_id: save_names_ages_sorted, points: 0, grade: false, solution: true -->
```python
def save_names_ages_sorted(names2ages, filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_save_names_ages_sorted, points: 10, grade: true, solution: false -->
```python
names2ages = {'Angela Merkel': 63,
              'Barack Obama': 56,
              'Barak Obama': 56,
              'Donald Trump': 71,
              'Emmanuel Macron': 39,
              'Justin Trudeau': 45,
              'Paul McCartney': 75,
              'Yoko Ono': 84}

save_names_ages_sorted(names2ages, 'new-names-ages.txt')

# a packace from Python's Standard Library for comparing two files
import filecmp
assert filecmp.cmp('names-ages.txt', 'new-names-ages.txt')

```